**Exercise**

Complete the following code to print the string `"hello"`.

<!-- grade_id: hello, points: 0, grade: false, solution: true -->
```python
def hello():
    ### BEGIN SOLUTION
    return 'hello'
    ### END SOLUTION
```

<!-- grade_id: test-hello, points: 1, grade: true, solution: false -->
```python
# tests
assert isinstance(hello(), str)
```

<!-- grade_id: htest-hello, points: 1, grade: true, solution: false -->
```python
# hidden tests
### BEGIN HIDDEN TESTS
assert hello() == 'hello'
### END HIDDEN TESTS
```

**Exercise**

What character does `"\n"` represents in python?

<!-- grade_id: character, points: 1, grade: true, solution: true -->
The newline character.